```python
# Initialize Otter
import otter
grader = otter.Notebook("mohrs_circle_notebook.ipynb")
```

# CEE 175: Geotechnical and Geoenvironmental Engineering

> **Alex Frantzis** <br> Cool Student >:3, UC Berkeley

[![License](https://img.shields.io/badge/license-CC%20BY--NC--ND%204.0-blue)](https://creativecommons.org/licenses/by-nc-nd/4.0/)
***

Assignment Description

```python
# Please run this cell, and do not modify the contents

import hashlib
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
    
import numpy as np
import math
import matplotlib.pyplot as plt
import matplotlib.image as mpimg
```

In this assignment we'll be calculating and plotting the Mohr's circle in order to find the modes of maximum stress. We'll use the mathematically derived equations of the Mohr's circle and verify our results with the graphical methods for Mohr's circle. As a warm up, let's start off with some basic equations using the stress equations.

## Question 1: Stress Calculation

Write three functions: `calcNormalX`, `calcNormalY`, and `calcShear` that calculate the normal and shear stresses on a plane at an arbitrary angle. The functions have the following input argument:
* `normalX`, which is the original normal stress in the x direction.
* `normalY`, which is the original normal stress in the y direction.
* `shear`, which is the original shear stress.
* `theta`, which is the angle of rotation in degrees.

and the following output arguments.
* `calcNormalX` outputs `new_normalX`, which is the normal stress in the x direction on the plane rotated counterclockwise at angle theta.
* `calcNormalX` outputs`new_normalY`, which is the normal stress in the y direction on the plane rotated counterclockwise at angle theta.
* `calcShear` outputs `new_shear`, which is the new shear stress on the plane rotated counterclockwise at angle theta.

####  Note: The input theta can be in either scalar or vector (np.array) form. You shouldn't have to think about this since we give you the correct functions to use to implement this trait, but it's good to think about when we use the function later.
---
```PYTHON
Examples:

>>> calcNormalX(55,25,20,0)
55

>>> calcNormalY(55,25,20,0)
25

>>> calcShear(55,25,20,0)
20

>>> calcNormalX(55,25,20,90)
25

>>> calcNormalY(55,25,20,90)
55

>>> calcShear(55,25,20,90)
20

>>> calcNormalX(200, 10, 100, 30)
239.1025

>>> calcNormalY(200, 10, 100, 30)
-29.1025

>>> calcShear(200, 10, 100, 30)
-32.2724
```
---
For a 2D stress state with normal stresses $\sigma_x, \sigma_y$ and shear stress $\tau_{xy}$,  
the **normal** and **shear** stresses on a plane rotated by angle $\theta$ (counterclockwise) are:

**Normal stress:**

$$
\sigma_{x'x'}
= \frac{\sigma_{xx}+\sigma_{yy}}{2}
+ \frac{\sigma_{xx}-\sigma_{yy}}{2}\cos(2\theta)
+ \tau_{xy}\sin(2\theta)
$$

$$
\sigma_{y'y'}
= \frac{\sigma_{xx}+\sigma_{yy}}{2}
- \frac{\sigma_{xx}-\sigma_{yy}}{2}\cos(2\theta)
- \tau_{xy}\sin(2\theta)
$$

**Shear stress:**

$$
\tau_{x'y'}
= -\frac{\sigma_x-\sigma_y}{2}\sin(2\theta)
+ \tau_{xy}\cos(2\theta)
$$

<center><img src="resources/Mohrs_Diagram.png" width='650'/></center>
<center>
- Stress Diagram
</center>

 
- $\theta$ is the physical rotation angle of the plane (radians or degrees).  
- Principal stresses occur where $\tau_\theta = 0$, and their values are $\sigma_{1,3} = C \pm R$.
- On Mohr's circle, a plane rotated by $\theta$ corresponds to a point at angle $2\theta$ from the positive $\sigma$-axis.
---
Note: We've already imported the numpy library. It is recommended that you use the trigonometric functions from there to implement this function. Be careful, since np.cos and np.sin take in radians as the default. We recommend querying an AI tool or looking at the documentation at https://numpy.org/doc/2.3/ to help in figuring out how to use these functions and their parameters.

```python
# ANSWER CELL

# We have seperated the function into 3 subfunctions for you to implement.
def calcNormalX(normalX, normalY, shear, theta):
    new_normalX = ...
    return new_normalX

def calcNormalY(normalX, normalY, shear, theta):
    newNormalY = ...
    return newNormalY

def calcShear(normalX, normalY, shear, theta):
    new_shear = ...
    return new_shear
```

```python
# ANSWER CELL

def calcStress(normalX, normalY, shear, theta):
    new_normalX = calcNormalX(normalX, normalY, shear, theta)
    new_normalY = calcNormalY(normalX, normalY, shear, theta)
    new_shear = calcShear(normalX, normalY, shear, theta)

    return new_normalX, new_normalY, new_shear
```

```python
# TEST YOUR FUNCTION HERE

q1 = ...

# print result
print(f'New Stresses = {q1} %')
np.shape(q1)
```

```python
grader.check("q1")
```

## Question 2: Optimize stresses

When dealing with stress diagrams, we often want to know what the worst-case scenario for the stresses is. We will explore a slightly different approach to doing this, opting to use an optimization algorithm on our function to determine the angle which maximizes the shear stress on a given stress distribution. Later on (or possibly already in your class), we will use the graphical method and see that this solution matches exactly in line with it.

---

An example of what the process should look like has been given below for you. We have imported the function `minimize_scalar` from the `scipy.optimize` library (https://docs.scipy.org/doc/scipy/reference/generated/scipy.optimize.minimize_scalar.html). Try looking at the documentation or querying AI to get an explanation of its functionality and parameters.


```python
from scipy.optimize import minimize_scalar

def maxNormalXAngle(normalX, normalY, shear):

    #Create dummy function that takes in theta and returns the negative value of the normal stress in the x direction
    def calculateNormalX(theta):
        new_normalX = calcNormalX(normalX, normalY, shear, theta)
        return -new_normalX

    # calculate the angle which minimizes -new_normalX, or equivalently maximizes newNormalX.
    optimization_result = minimize_scalar(calculateNormalX, bounds = (0, 180), method = 'bounded')

    optimal_theta = optimization_result.x
    max_normalX, _, _ = calcStress(normalX, normalY, shear, optimal_theta)
    return optimal_theta, max_normalX

#Testing function
q2_test = maxNormalXAngle(67, 67, 21)
print(f'\noptimal_theta = {round(q2_test[0], 1)}°')
print(f'\nmax normal stress = {round(q2_test[1], 1)}')
```

Write a function named `maxShearAngle()` that returns the angle in which the returned shear stress is highest. This function has the following input argument:
* `normalX`, which is the original normal stress in the x direction.
* `normalY`, which is the original normal stress in the y direction.
* `shear`, which is the original shear stress.
* `theta`, which is the angle of rotation in degrees.

and the following output arguments.
* `optimal_theta`, which is the theta which maximizes the shear.
---
```PYTHON
Examples:

>>> maxShearAngle(30, 30, 40)
0.0

>>> maxShearAngle(100, 0, 0)
135.0

>>> maxShearAngle(0, 100, 0)
45.0

>>> maxShearAngle(80, -20, 20)
145.9
```
---
Hint: Try copy and pasting from the example and changing the values that you're optimizing on!

```python
# ANSWER CELL

def maxShearAngle(normalX, normalY, shear):

    #Create dummy function that takes in theta and returns the negative value of the normal stress in the x direction
    ...

    # calculate the angle which minimizes -new_normalX, or equivalently maximizes newNormalX.
    ...
    
    return optimal_theta, max_shear
```

```python
# TEST YOUR FUNCTION HERE
q2 = ...

# print result
print(f'optimal_theta = {round(q2[0], 1)}°')
print(f'\nmax shear stress = {round(q2[1], 1)}')
```

```python
grader.check("q2")
```

## Question 3: Plotting

Now we will use our function from question 0 to plot the mohrs circle diagram and compare the results to what we expect and what we have done so far. The stress state will be 
* `normalX` = 1500
* `normalY` 200
* `shear` = 100

---
You will first use the `calcStress(normalX, normalY, shear, theta)` function you created previously to calculate the normalX and shear values corresponding to each angle in the array `theta_vals` that we have created for you.

Then you will create a `matplotlib.pyplot` figure stored in `fig_1` and perform the following tasks on the same figure:
* Plot the x axis normal stress values vs the shear values with normalX_vals on the X-axis and shear_vals on the Y-axis. 
* Set the figure title to  "Mohr circle diagram".
* Set the x-axis label to "normal stress, MPa".
* Set the y-axis label to "shear stress, MPa"
* Set the x-axis limits to (0,1600).
* Set the y-axis limits to (-800,800).
---
If you're running into trouble or have not used matplotlib before, check the documentation here (https://matplotlib.org/stable/users/index.html) or again, use AI as a tool to see what each function does and how they work together.

```python
# ANSWER CELL
import matplotlib.pyplot as plt

# Plot inside this function for grading purposes
def makeFigure():
    # Create figure. Do not modify this line for grading purposes
    fig_1 = plt.figure(figsize=(8,8))
    
    # array of values for theta, ranging from [0, 180] with 360 equally spaced elements
    theta_vals = np.linspace(0,180,360)
    
    # Calculate stress values
    normalX_vals, normalY_vals, shear_vals = ...
    
    # Plot mohrs circle using stress values
    plt.plot(...)
    
    # Set figure title
    plt.title(...)
    
    # Set axes labels
    plt.xlabel(...)
    plt.ylabel(...)
    
    # Set axis limits 
    plt.xlim(...)
    plt.ylim(...)
    
    # Add gridlines
    plt.grid()
    
    return fig_1

fig_1 = makeFigure()
```

```python
grader.check("q3")
```

## Hooray! You finished the graded portion :D
For your convenience, we have added a function below implementing the pole method that you learned in class. Try comparing it with what we've made so far to see how the results match up!

```python
def plot_mohr_circle(sigma_h, tau_h, sigma_v, tau_v):
    """
    Calculates and plots Mohr's Circle for a given 2D stress state
    using the geotechnical sign convention.

    Args:
        sigma_h (float): Normal stress on the horizontal plane (+ comp).
        tau_h (float): Shear stress on the horizontal plane (+ ccw).
        sigma_v (float): Normal stress on the vertical plane (+ comp).
        tau_v (float): Shear stress on the vertical plane (+ ccw).
    """
    # --- 1. Calculate Mohr's Circle Parameters ---
    C = (sigma_h + sigma_v) / 2
    R = np.sqrt(((sigma_h - sigma_v) / 2)**2 + tau_h**2)

    print(f"Center of the Circle (C): ({C:.2f}, 0)")
    print(f"Radius of the Circle (R): {R:.2f}")

    # --- 2. Calculate Key Stress Values ---
    sigma_1 = C + R  # Major Principal Stress (max compression)
    sigma_3 = C - R  # Minor Principal Stress (max tension)
    tau_max = R
    sigma_at_tau_max = C

    print(f"\nMajor Principal Stress (sigma_1): {sigma_1:.2f} kPa")
    print(f"Minor Principal Stress (sigma_3): {sigma_3:.2f} kPa")
    print(f"Maximum Shear Stress (tau_max): {tau_max:.2f} kPa")

    # --- 3. Determine the Pole ---
    # Intersection of a horizontal line from H and a vertical line from V
    pole = (sigma_v, tau_h)
    print(f"Coordinates of the Pole (P): {pole}")

    # --- 4. Plotting ---
    fig, ax = plt.subplots(figsize=(10, 10))
    ax.set_aspect('equal')
    ax.grid(True, which='both', linestyle='--', linewidth=0.5)

    # Create and plot the circle
    angles = np.linspace(0, 2 * np.pi, 100)
    sigma_circle = C + R * np.cos(angles)
    tau_circle = R * np.sin(angles)
    ax.plot(sigma_circle, tau_circle, 'b-', label="Mohr's Circle")

    # Plot axes
    ax.axhline(y=0, color='k')
    ax.axvline(x=0, color='k', linestyle=':')

    # Plot key points
    ax.plot(C, 0, 'ko', markersize=8, label=f'Center C ({C:.1f}, 0)')
    ax.plot(sigma_h, tau_h, 'ro', markersize=8, label=f'Horizontal Plane H ({sigma_h}, {tau_h})')
    ax.plot(sigma_v, tau_v, 'go', markersize=8, label=f'Vertical Plane V ({sigma_v}, {tau_v})')
    ax.plot(pole[0], pole[1], 'm*', markersize=12, label=f'Pole P {pole}')
    ax.plot(sigma_1, 0, 'yo', markersize=8, label=f'$\\sigma_1$ = {sigma_1:.1f}')
    ax.plot(sigma_3, 0, 'yo', markersize=8, label=f'$\\sigma_3$ = {sigma_3:.1f}')
    ax.plot(sigma_at_tau_max, tau_max, 'cs', markersize=8, label=f'$\\tau_{{max}}$ = {tau_max:.1f}')
    ax.plot(sigma_at_tau_max, -tau_max, 'cs', markersize=8)

    # Plot connecting lines
    ax.plot([sigma_h, sigma_v], [tau_h, tau_v], 'k--') # Diameter
    ax.plot([pole[0], sigma_1], [pole[1], 0], 'y--', label='Major Principal Plane')

    # Formatting
    ax.set_title("Mohr's Circle (Geotechnical Convention)", fontsize=16)
    ax.set_xlabel('Normal Stress, $\\sigma$ (kPa) | (+) Compression', fontsize=12)
    ax.set_ylabel('Shear Stress, $\\tau$ (kPa) | (+) CCW', fontsize=12)
    ax.legend(loc='best', bbox_to_anchor=(1.05, 1))

    # Invert y-axis for standard geotechnical plot representation
    plt.gca().invert_yaxis()

    plt.show()


if __name__ == '__main__':
    # --- Input Stresses (Geotechnical Convention) from the screenshot ---
    # Positive (+) for Compression, Negative (-) for Tension
    # Positive (+) for Counter-Clockwise (CCW) shear, Negative (-) for Clockwise (CW)

    # Stresses on the Horizontal (H) Plane:
    sigma_horizontal = 100  # 100 kPa Compression -> Positive
    tau_horizontal = -35    # 35 kPa Clockwise shear -> Negative

    # Stresses on the Vertical (V) Plane:
    sigma_vertical = -20  # 20 kPa Tension -> Negative
    tau_vertical = 35     # 35 kPa Counter-Clockwise shear -> Positive

    plot_mohr_circle(sigma_horizontal, tau_horizontal, sigma_vertical, tau_vertical)
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 0)
```
>>> assert get_hash(type(q1)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> assert np.isclose(calcStress(55, 25, 20, 0), [55, 25, 20], rtol=0.01, atol=0.01).all() == True
>>> assert np.isclose(calcStress(55, 25, 20, 90), [25, 55, -20], rtol=0.01, atol=0.01).all() == True
>>> assert np.isclose(calcStress(200, 10, 100, 30), [239.1025, -29.1025, -32.2724], rtol=0.1, atol=0.1).all() == True
>>> assert get_hash(int(calcStress(80, 0, 10, 45)[0])) == 'c0c7c76d30bd3dcaefc96f40275bdc0a'
>>> assert get_hash(int(calcStress(80, 0, 10, 45)[1])) == '34173cb38f07f89ddbebc2ac9128303f'
>>> assert get_hash(int(calcStress(80, 0, 10, 45)[2])) == '046be5694dce5d772d74b8f6964149a2'
>>> assert get_hash(int(calcStress(0, 0, 50, 45)[0])) == 'c0c7c76d30bd3dcaefc96f40275bdc0a'
>>> assert get_hash(int(calcStress(0, 0, 50, 45)[1])) == '25d3122a93ce9214a42700a394acb5f4'
>>> assert get_hash(int(calcStress(0, 0, 50, 45)[2])) == 'cfcd208495d565ef66e7dff9f98764da'
>>> assert get_hash(int(calcStress(55, 25, 20, 180)[0])) == 'a684eceee76fc522773286a895bc8436'
>>> assert get_hash(int(calcStress(55, 25, 20, 180)[1])) == '8e296a067a37563370ded05f5a3bf3ec'
>>> assert get_hash(int(calcStress(55, 25, 20, 180)[2])) == '98f13708210194c475687be6106a3b84'
>>> assert get_hash(int(calcStress(30, 30, 40, 60)[0])) == 'ea5d2f1c4608232e07d3aa3d998e5135'
>>> assert get_hash(int(calcStress(30, 30, 40, 60)[1])) == '0267aaf632e87a63288a08331f22c7c3'
>>> assert get_hash(int(calcStress(30, 30, 40, 60)[2])) == 'b9ea889c6fc3fb2b4c343d7400734856'
>>> assert get_hash(int(calcStress(0, 100, 0, 30)[0])) == '1ff1de774005f8da13f42943881c655f'
>>> assert get_hash(int(calcStress(0, 100, 0, 30)[1])) == 'd09bf41544a3365a46c9077ebb5e35c3'
>>> assert get_hash(int(calcStress(0, 100, 0, 30)[2])) == '17e62166fc8586dfa4d1bc0e1742c08b'
>>> assert get_hash(int(calcStress(0, 100, 0, 60)[0])) == 'ad61ab143223efbc24c7d2583be69251'
>>> assert get_hash(int(calcStress(0, 100, 0, 60)[1])) == '8e296a067a37563370ded05f5a3bf3ec'
>>> assert get_hash(int(calcStress(0, 100, 0, 60)[2])) == '17e62166fc8586dfa4d1bc0e1742c08b'
```
failure_message: Your funtion is getting some calculations wrong!
success_message: Function correctly classifies all scalar test cases :D
```
>>> assert get_hash(int(np.sum(calcStress(1500, 200, 100, np.linspace(0, 180, 360))[0]))) == '51972cf4492f4879a16614c849c01da7'
>>> assert get_hash(int(np.sum(calcStress(1500, 200, 100, np.linspace(0, 180, 360))[1]))) == '582eff45b7eb4e3a3484ba5d72ac4f87'
>>> assert get_hash(int(np.sum(calcStress(1500, 200, 100, np.linspace(0, 180, 360))[2]))) == 'f899139df5e1059396431415e770c6dd'
>>> assert get_hash(int(np.mean(calcStress(1500, 200, 100, np.linspace(0, 180, 360))[0]))) == '92fb0c6d1758261f10d052e6e2c1123c'
>>> assert get_hash(int(np.mean(calcStress(1500, 200, 100, np.linspace(0, 180, 360))[1]))) == '362e80d4df43b03ae6d3f8540cd63626'
>>> assert get_hash(int(np.mean(calcStress(1500, 200, 100, np.linspace(0, 180, 360))[2]))) == 'cfcd208495d565ef66e7dff9f98764da'
>>> assert get_hash(int(calcStress(1500, 200, 100, np.linspace(0, 180, 3))[0][0])) == 'cfa5301358b9fcbe7aa45b1ceea088c6'
>>> assert get_hash(int(calcStress(1500, 200, 100, np.linspace(0, 180, 3))[0][1])) == '3644a684f98ea8fe223c713b77189a77'
>>> assert get_hash(int(calcStress(1500, 200, 100, np.linspace(0, 180, 3))[0][2])) == 'cfa5301358b9fcbe7aa45b1ceea088c6'
>>> assert get_hash(int(calcStress(1500, 200, 100, np.linspace(0, 180, 3))[1][0])) == '3644a684f98ea8fe223c713b77189a77'
>>> assert get_hash(int(calcStress(1500, 200, 100, np.linspace(0, 180, 3))[1][1])) == 'cfa5301358b9fcbe7aa45b1ceea088c6'
>>> assert get_hash(int(calcStress(1500, 200, 100, np.linspace(0, 180, 3))[1][2])) == '3644a684f98ea8fe223c713b77189a77'
>>> assert get_hash(int(calcStress(1500, 200, 100, np.linspace(0, 180, 3))[2][2])) == 'f899139df5e1059396431415e770c6dd'
```
failure_message: Your funtion is not behaving well on array inputs!
success_message: Function correctly classifies all array test cases :D


Tests for q2 (points: 0)
```
>>> assert get_hash(type(q2)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> assert get_hash(int(maxShearAngle(30, 30, 40)[1])) == 'd67d8ab4f4c10bf22aa353e27879133c'
>>> assert get_hash(int(maxShearAngle(100, 0, 0)[1])) == 'c0c7c76d30bd3dcaefc96f40275bdc0a'
>>> assert get_hash(int(maxShearAngle(0, 100, 0)[1])) == 'c0c7c76d30bd3dcaefc96f40275bdc0a'
>>> assert get_hash(int(maxShearAngle(200, 10, 100)[1])) == '3988c7f88ebcb58c6ce932b957b6f332'
>>> assert get_hash(int(maxShearAngle(50, 50, -25)[1])) == '8e296a067a37563370ded05f5a3bf3ec'
>>> assert get_hash(int(maxShearAngle(80, -20, 20)[1])) == 'd82c8d1619ad8176d665453cfb2e55f0'
```
failure_message: Your funtion is incorrectly calculating the max shear stress
success_message: Function correctly calculates max shear stress :D
```
>>> assert get_hash(int(maxShearAngle(30, 30, 40)[0])) == 'cfcd208495d565ef66e7dff9f98764da'
>>> assert get_hash(int(maxShearAngle(100, 0, 0)[0])) == '7f1de29e6da19d22b51c68001e7e0e54'
>>> assert get_hash(int(maxShearAngle(0, 100, 0)[0])) == 'f7177163c833dff4b38fc8d2872f1ec6'
>>> assert get_hash(int(maxShearAngle(200, 10, 100)[0])) == '06409663226af2f3114485aa4e0a23b4'
>>> assert get_hash(int(maxShearAngle(50, 50, -25)[0])) == '8613985ec49eb8f757ae6439e879bb2a'
>>> assert get_hash(int(maxShearAngle(80, -20, 20)[0])) == '2b24d495052a8ce66358eb576b8912c8'
```
failure_message: Your funtion is incorrectly calculating the optimum angle!
success_message: Function correctly calculates optimum angle :D


Tests for q3 (points: 0)
```
>>> fig_1 = makeFigure()
>>> assert type(fig_1).__name__ == 'Figure'
<Figure size 800x800 with 1 Axes>
```
failure_message: Make sure to create a figure and assign it to fig_1.
```
>>> fig_1 = makeFigure()
>>> assert get_hash(np.round(np.sum(fig_1.axes[0].get_lines()[0].get_xdata()), decimals=2)) == '47f6d46a79d004de38c6d3c1174ca637'
>>> assert get_hash(np.round(np.sum(fig_1.axes[0].get_lines()[0].get_ydata()), decimals=2)) == '62dca49f0781bf26b4305bddb0414bea'
<Figure size 800x800 with 1 Axes>
```
failure_message: Check x and y data of plot.
```
>>> fig_1 = makeFigure()
>>> x_string = ['NORMAL', 'STRESS', 'MPA']
>>> assert all([word in fig_1.axes[0].get_xlabel().upper() for word in x_string]), 'Check the x-axis label.'
>>> y_string = ['SHEAR', 'STRESS', 'MPA']
>>> assert all([word in fig_1.axes[0].get_ylabel().upper() for word in y_string]), 'Check the y-axis label.'
<Figure size 800x800 with 1 Axes>
```
failure_message: Check figure labels
```
>>> fig_1 = makeFigure()
>>> title_string = ['CIRCLE', 'DIAGRAM']
>>> assert all([word in fig_1.axes[0].get_title().upper() for word in title_string]), 'Check the title label.'
<Figure size 800x800 with 1 Axes>
```
failure_message: Check title.
```
>>> fig_1 = makeFigure()
>>> assert fig_1.axes[0].get_xlim() == (0, 1600), 'Check the x-axis limits.'
>>> assert fig_1.axes[0].get_ylim() == (-800, 800), 'Check the y-axis limits.'
<Figure size 800x800 with 1 Axes>
```
failure_message: Check axis limits.
