```python
# Initialize Otter
import otter
grader = otter.Notebook("Effective_Stress_notebook.ipynb")
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("Effective_Stress_notebook.ipynb")
```

# CEE 175: Geotechnical and Geoenvironmental Engineering

> **Alex Frantzis** <br> Cool Student >:3, UC Berkeley

[![License](https://img.shields.io/badge/license-CC%20BY--NC--ND%204.0-blue)](https://creativecommons.org/licenses/by-nc-nd/4.0/)
***

## Understanding Total Stress, Pore Pressure, and Effective Stress

In soil mechanics, the **total stress**, **pore water pressure**, and **effective stress** are key concepts that describe how forces are distributed within the ground.

### Total Stress (σ)
Total stress is the **entire weight of the soil and any water above a certain depth**.
It represents the pressure from the overlying soil layers acting downward due to gravity.

$$
\sigma = \gamma \, z
$$

where:  
$$
\begin{aligned}
\text{where:} \\
\sigma & = \text{total stress (kPa)} \\
\gamma & = \text{unit weight of the soil (kN/m³)} \\
z & = \text{depth below the surface (m)}
\end{aligned}
$$



---

### Pore Water Pressure (u)
Pore water pressure is the **pressure exerted by water** in the pores of the soil.
It only exists **below the water table**, where the soil is fully saturated.

$$
u = \gamma_w \, (z - z_w)
$$

$$
\begin{aligned}
\text{where:} \\
u & = \text{pore water pressure (kPa)} \\
\gamma_w & = \text{unit weight of water (≈ 9.81 kN/m³)} \\
z_w & = \text{depth to the water table (m)} \\
\text{For } z < z_w, & \; u = 0 \; \text{(no pore pressure above the water table)}
\end{aligned}
$$


---

### Effective Stress (σ′)
Effective stress is the **stress carried by the soil skeleton** — it determines how the soil actually behaves (strength, deformation, and stability).

$$
\sigma' = \sigma - u
$$

$$
\begin{aligned}
\text{where:} \\
\sigma' & = \text{effective stress (kPa)} \\
\sigma & = \text{total stress (kPa)} \\
u & = \text{pore water pressure (kPa)}
\end{aligned}
$$



\begin{aligned}
\textbf{This means:} \\
\text{Above the water table:} & \quad u = 0 \;\Rightarrow\; \sigma' = \sigma \\
\text{Below the water table:} & \quad \text{part of the total stress is supported by water, reducing } \sigma'
\end{aligned}



---

### Key Idea
Think of the soil as a sponge:
- The **total stress** is the weight of everything pressing down.  
- The **pore pressure** is the pressure of water filling the pores.  


```python
# Please run this cell, and do not modify the contents

import hashlib
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()

import pandas as pd
import numpy as np
import math
import matplotlib.pyplot as plt
import matplotlib.image as mpimg
```

## Question 0: Total Stress Calculation

Write a function named `getStress()` that calculates the total vertical stress (σ) at specified depth. This function has the following input argument:
* `thickness`, the depth of the point where stress is to be determined
* `thickness`, the thicknesses of the soil layers, arranged from the surface layer (closest to ground) to the deepest layer below ground.
* `density`, the unit weights of the respective soil layers, in kN/m³.

and the following output arguments.
* `stress`, the stress at specified depth.
---
```PYTHON
Example:

    thicknesses = [0.5, 1.0, 2.0, 1.5, 3.0]
    densities = [17.5, 18.0, 19.0, 20.5, 21.0]

>>> stressAtDepth(2, thicknesses, densities)
36.25

>>> stressAtDepth(2, thicknesses, densities)
75.0

>>> stressAtDepth(2, thicknesses, densities)
158.5
```
---
The following is an example of the inputs to the function. Make sure to run it so that the given testing cell has something to test on. For reference on pandas check https://pandas.pydata.org/docs/reference/index.html

```python
# Define soil layers by their thickness and density
layers = {
    "thickness": [0.5, 1.0, 2.0, 1.5, 3.0],   # layer thicknesses in meters
    "density": [17.5, 18.0, 19.0, 20.5, 21.0] # unit weights (kN/m³)
}

df = pd.DataFrame(layers)

# Compute cumulative depth at the bottom of each layer
df["depth"] = df["thickness"].cumsum()

print(df)
```

```python

def getStress(z, thickness, density):
    stress = ...
    remaining_depth = ...
        
    for i in range(len(thickness)):
        layer_thickness = ...
        layer_density = ...

        # If depth is within this layer
        if remaining_depth <= layer_thickness:
            ...
            break
        else:
            # Full layer contributes its entire weight
            ...
            ...
    return stress
    
                

```

```python
# TEST YOUR FUNCTION HERE
depth = ...
q0 = ...

# print result
print(f'Stress at depth {depth} = {q0} kPa')
```

```python
grader.check("q0")
```

```python
def getPP(depth, water_table_depth, gamma_w=9.81):
    if depth <= water_table_depth:
        ...
    else:
        # Pressure increases linearly below the water table
        ...

    return porePressure
```

```python
# TEST YOUR FUNCTION HERE
depth = ...
waterTable = 2
q1 = getPP(depth, waterTable)

# print result
print(f'Stress at depth {depth} with water table at height {waterTable} is equal to {q1} kPa')
```

```python
grader.check("q1")
```

## Question 2: Effective Stress Calculation

Write a function named `getEffStress()` which takes a given depth and calculates the effective stress (σ') at that specified depth. This function uses the total stress from the soil layers and subtracts the pore water pressure below the water table. It has the following input arguments.
* `thickness`, which are the thicknesses of the soil layers arranged in order from closest to ground to lowest below ground.
* `density`, which are the unit weights (kN/m³) of the respective layers.
* `water_table_depth`, which is the depth (in meters) from the surface to the groundwater table.
* (optional) `gamma_w`, the unit weight of water (kN/m³). Default is 9.81.

and the following output arguments.
* `effStress`, the effective stress (kPa or kN/m²) at that depth.
---
    Example:
```PYTHON

    thicknesses = [0.5, 1.0, 2.0, 1.5, 3.0]
    densities = [17.5, 18.0, 19.0, 20.5, 21.0]

>>> getEffStress(1, thicknesses, densities, water_table_depth=2.0)
17.5

>>> getEffStress(4, thicknesses, densities, water_table_depth=2.0)
75

>>> getEffStress(7, thicknesses, densities, water_table_depth=2.0)
137.5
```

```python
def getEffStress(depth, thickness, density, water_table_depth, gamma_w=9.81):
    
    # Get total stress and pore pressure
    totalStress = ...
    porePressure = ...

    ...

```

```python
# TEST YOUR FUNCTION HERE
depth = ...
q2 = ...

# print result
print(f'Stress at depth {depth} = {q2} kPa')
```

```python
grader.check("q2")
```

## Question 3: Plotting Total, Pore, and Effective Stress vs. Depth

In this section, you will **visualize the stress distribution** in the soil by plotting the following quantities as functions of depth:

- **Total Stress** — the total vertical stress due to the weight of the overlying soil.  
- **Pore Water Pressure** — the pressure exerted by water below the water table.  
- **Effective Stress** — the portion of the stress carried by the soil skeleton, defined as:

$$
\sigma' = \sigma - u
$$

---

### Instructions

1. Use the previously defined functions:
   - `getStress()`
   - `getPP()`
   - `getEffStress()`

2. Define your soil profile:
   - Layer **thicknesses** and **densities**
   - The **water table depth**

3. Compute:
   - Total stress 
   - Pore water pressure 
   - Effective stress 

   for a range of depths (z) (e.g., from 0 m to the total soil depth in small increments).

4. Plot **Total Stress**, **Pore Pressure**, and **Effective Stress** on the same graph.

---

### Example Plot Layout

- **x-axis:** Stress (kPa) 
- **y-axis:** Depth (m) — increasing **downward**

Use different line styles or colors to distinguish:<br>
— Total Stress  
— Pore Pressure  
— Effective Stress  

Make sure to include:
- A **legend** identifying each curve  
- **Axis labels** with units  
- An **inverted y-axis** so that depth increases downward, matching geotechnical convention.
--- 
### Use the parameters as discerned from the figure below

```python
# -----------------------------
# Define soil layers and water table
# -----------------------------

thickness = ...
densities = ...
water_table_depth = ...
gamma_w = 9.81           # kN/m^3

# -----------------------------
# Sample depths and compute values
# -----------------------------
depth = np.cumsum(thickness) # Adds up thicknesses. Last value is the max depth.
z_max = depth[-1]
z = np.linspace(0, z_max, 400) # Ad
 
totalStress = np.array([getStress(zi, thickness, densities)     for zi in z])
porePressure     = ...
effStress = ...

# -----------------------------
# Plot
# -----------------------------
fig_1 = plt.figure(figsize=(6, 7)) # Do not change this line

plt.plot(totalStress,   z, label=r"Total Stress $\sigma(z)$") 
...
...

# Water table reference line
plt.axhline(y=water_table_depth, linestyle="--", linewidth=1, label="Water Table")

# Axes formatting
plt.gca().invert_yaxis()  # depth increases downward
...
...
...
plt.legend()
plt.grid(True, linestyle=":", linewidth=0.8)
plt.tight_layout()
plt.show()

```

```python
grader.check("q3")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q0 (points: 0)
```
>>> assert get_hash(type(q0)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> layers = {'thickness': [0.5, 1.0, 2.0, 1.5, 3.0], 'density': [17.5, 18.0, 19.0, 20.5, 21.0]}
>>> df = pd.DataFrame(layers)
>>> assert get_hash(int(10 * np.round(getStress(2, df['thickness'], df['density']), 1))) == 'c3e878e27f52e2a57ace4d9a76fd9acf'
>>> assert get_hash(int(10 * np.round(getStress(4, df['thickness'], df['density']), 1))) == 'b137fdd1f79d56c7edf3365fea7520f2'
>>> assert get_hash(int(10 * np.round(getStress(8, df['thickness'], df['density']), 1))) == '228499b55310264a8ea0e27b6e7c6ab6'
```
failure_message: Your funtion is getting some calculations wrong!
success_message: Function correctly calculates stress for all depth test cases :D
```
>>> layers = {'thickness': [0.5, 1.0, 2.0, 1.5, 3.0], 'density': [17.5, 18.0, 19.0, 20.5, 21.0]}
>>> df = pd.DataFrame(layers)
>>> assert get_hash(int(10 * np.round(getStress(5, df['thickness'], df['density']), 1))) == 'ef4e3b775c934dada217712d76f3d51f'
>>> assert get_hash(int(10 * np.round(getStress(10, df['thickness'], df['density']), 1))) == '228499b55310264a8ea0e27b6e7c6ab6'
>>> assert get_hash(int(10 * np.round(getStress(20, df['thickness'], df['density']), 1))) == '228499b55310264a8ea0e27b6e7c6ab6'
```
failure_message: Your funtion is getting some calculations wrong!
success_message: Function correctly calculates stress for all depth test cases :D


Tests for q1 (points: 0)
```
>>> assert get_hash(type(q1)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> assert get_hash(int(10 * np.round(getPP(1, water_table_depth=2), 1))) == 'cfcd208495d565ef66e7dff9f98764da'
>>> assert get_hash(int(10 * np.round(getPP(3, water_table_depth=2), 1))) == 'ed3d2c21991e3bef5e069713af9fa6ca'
>>> assert get_hash(int(10 * np.round(getPP(7, water_table_depth=2), 1))) == '559cb990c9dffd8675f6bc2186971dc2'
```
failure_message: Your function is getting some calculations wrong!
success_message: Function correctly calculates pore pressure for all depth test cases :D


Tests for q2 (points: 0)
```
>>> assert get_hash(type(q1)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> layers = {'thickness': [2, 5, 7, 10], 'density': [100, 200, 150, 230]}
>>> df = pd.DataFrame(layers)
>>> assert get_hash(int(10 * np.round(getEffStress(1, df['thickness'], df['density'], 4), 1))) == 'a9b7ba70783b617e9998dc4dd82eb3c5'
>>> assert get_hash(int(10 * np.round(getEffStress(7, df['thickness'], df['density'], 4), 1))) == 'f2b1b9275b094c26ff1998a09cd9412d'
>>> assert get_hash(int(10 * np.round(getEffStress(20, df['thickness'], df['density'], 4), 1))) == '9b3e5c1c0754bc6a379163afabe2af79'
```
failure_message: Your funtion is getting some calculations wrong!
success_message: Function correctly calculates stress for all depth test cases :D
```
>>> layers = {'thickness': [0.5, 1.0, 2.0, 1.5, 3.0], 'density': [17.5, 18.0, 19.0, 20.5, 21.0]}
>>> df = pd.DataFrame(layers)
>>> assert get_hash(int(10 * np.round(getEffStress(1, df['thickness'], df['density'], 2), 1))) == '8f85517967795eeef66c225f7883bdcb'
>>> assert get_hash(int(10 * np.round(getEffStress(4, df['thickness'], df['density'], 2), 1))) == '5e388103a391daabe3de1d76a6739ccd'
>>> assert get_hash(int(10 * np.round(getEffStress(7, df['thickness'], df['density'], 2), 1))) == '170c944978496731ba71f34c25826a34'
```
failure_message: Your funtion is getting some calculations wrong!
success_message: Function correctly calculates stress for all depth test cases :D


Tests for q3 (points: 0)
```
>>> assert get_hash(np.round(np.sum(fig_1.axes[0].get_lines()[0].get_xdata()), decimals=0)) == '8ef1bc3e5e466b6c7c905f19ae690f7a'
>>> assert get_hash(np.round(np.sum(fig_1.axes[0].get_lines()[0].get_ydata()), decimals=0)) == '6febd8047eb9ddc4d9caf141403a586a'
```
failure_message: Check x and y data of line plot 1.
```
>>> assert get_hash(np.round(np.sum(fig_1.axes[0].get_lines()[1].get_xdata()), decimals=0)) == 'f9f0ef796b1fc62af121ece66023b801'
>>> assert get_hash(np.round(np.sum(fig_1.axes[0].get_lines()[1].get_ydata()), decimals=0)) == '6febd8047eb9ddc4d9caf141403a586a'
```
failure_message: Check x and y data of line plot 1.
```
>>> assert get_hash(np.round(np.sum(fig_1.axes[0].get_lines()[2].get_xdata()), decimals=0)) == '3419bd55372df1e7f323ae0c62153ee7'
>>> assert get_hash(np.round(np.sum(fig_1.axes[0].get_lines()[2].get_ydata()), decimals=0)) == '6febd8047eb9ddc4d9caf141403a586a'
```
failure_message: Check x and y data of line plot 1.
```
>>> x_string = ['STRESS', 'KPA']
>>> assert all([word in fig_1.axes[0].get_xlabel().upper() for word in x_string]), 'Check the x-axis label.'
>>> y_string = ['DEPTH', 'M']
>>> assert all([word in fig_1.axes[0].get_ylabel().upper() for word in y_string]), 'Check the y-axis label.'
```
failure_message: Check figure labels and font sizes.
